# Grundlagen der Künstlichen Intelligenz - Wintersemester 2024/25

# Übung 7: Reinforcement Learning

---

> 'Grundlagen der künstlichen Intelligenz' im Wintersemester 2024/2025
>
> - T.T.-Prof. Benjamin Schäfer, benjamin.schaefer@kit.edu
> - Prof. Gerhard Neumann, gerhard.neumann@kit.edu

---

In dieser Übung werden wir ein approximatives Q-Lernen mit linearer Funktionsannäherung. Dafür implementieren wir die Berechnung von Q-Werten, Softmax-Policy zur Aktionswahl, den Temporal Difference (TD)-Fehler sowie die Parameteraktualisierung mittels Gradientenabstieg.

### Übungsteam

- Philipp Dahlinger, philipp.dahlinger@kit.edu
- Nicolas Schreiber, nicolas.schreiber@kit.edu
- Sebastian Pütz, sebastian.puetz@kit.edu
- Ulrich Oberhofer, ulrich.oberhofer@kit.edu


# Gruppenabgabe

Die Übungsblätter können in Gruppen von bis zu **3 Studierenden** abgegeben werden. **Jede Person aus der Gruppe muss die finale Version der Abgabe über Ilias hochladen**, es genügt nicht, dass nur eine Person aus der Gruppe dies tut. Es ist prinzipiell möglich, im Laufe des Semesters sich einer neuen Gruppe anzuschließen, sollte sich die eigene Gruppe vorzeitig auflösen. Generell muss jede Gruppe ihre eigene Lösung hochladen, wir werden die Abgaben auf Duplikate überprüfen.

Die Gruppen werden automatisch erfasst, **gebt deshalb die u-Kürzel eurer Gruppenmitglieder in die folgende Zelle ein.** Falls eure Gruppe nur aus 2 Studierenden besteht, oder ihr alleine abgibt, lasst die verbleibenden Felder frei. Hier ein Beispiel für eine Gruppe bestehend aus uabcd und uefgh:

_U-Kürzel der Gruppenmitglieder:_

_Mitglied 1: uabcd_

_Mitglied 2: uefgh_

_Mitglied 3:_


U-Kürzel der Gruppenmitglieder:

Mitglied 1:

Mitglied 2:

Mitglied 3:


# Auto-grading

Wir nutzen ein auto-grading System, welches eure abgegebenen Jupyter Notebooks automatisch analysiert und über
hidden Tests auf Richtigkeit prüft. Über diese Tests werden die Punkte bestimmt, die ihr für das Übungsblatt erhaltet.

Damit das auto-grading reibungslos funktioniert bitte folgende Dinge beachten:

- Vor dem Abgeben eines Notebooks bitte testen, dass alles von vorne bis hinten ohne Fehler durchläuft.
- Zellen, welche mit "### DO NOT CHANGE ###" markiert sind dürfen weder gelöscht noch bearbeitet werden
- Eure Lösung muss in die richtige Zelle (markiert mit "# YOUR CODE HERE") eingetragen werden.
    - (dabei natürlich den NotImplementedError löschen!)
- Es gibt potentiell scheinbar leere Zellen, die auch mit "### DO NOT CHANGE ###" markiert sind. Auch diese dürfen nicht bearbeitet oder gelöscht werden.
    - Falls dies doch gemacht wird, dann wird das automatische Grading nicht funktionieren und ihr erhaltet keine Punkte.
    - Wir werden hier strikt handeln und keine Ausnahmen machen, falls jemand doch Zellen verändert, die eindeutig als readonly markiert sind!
- Die Jupyter Notebooks haben inline Tests (für euch sichtbar), welche euer Ergebnis auf grobe Richtigkeit überprüfen.
    - Diese sind primär für euch, um Fehler zu erkennen und zu korrigieren.
    - Die inline Tests, die ihr im Notebook sehen könnt, sind allerdings nicht die Tests welche für das Grading verwendet werden!
    - Die inline Tests sind eine notwendige Bedingung, um beim Grading der Aufgabe Punkte zu erhalten!

# **WICHTIG** Abgabe des Notebooks
- Bitte das Jupyter Notebook mit dem ursprünglichen Dateinamen ins Ilias hochladen ("ex_06_hmm_mdp.ipynb")
- Bitte Jupyter Notebook und handgeschriebene PDF einzeln hochladen, nicht als ZIP.
- Bitte darauf achten, dass die Jupyter Notebook Zell-Metadaten erhalten bleiben. Das ist eigentlich immer der Fall,
in wenigen Fällen gab es hier jedoch Probleme. Um auf Nummer Sicher zu gehen bitte das Notebook vor der Abgabe ein Mal
in einem normalen Texteditor öffnen und nach "nbgrader" suchen. Wenn hier dann keine entsprechenden JSON-Einträge auftauchen
dann sind leider die Metadaten verloren gegangen.

# Aufgabe 2: SnaQe
Diese Übung wurde von Philipp Dahlinger für die KIT-Vorlesung Grundlagen der Künstlicher Intelligenz. Die pygame-Implementierung der Snake-Umgebung wurde aus dieser Quelle adaptiert: [GitHub](https://github.com/jl4r1991/SnakeQlearning)

## Anweisungen

**Führe jeden Codeblock aus, um keine Dependency zu verpassen.**

Dieses Jupyter-Notebook bietet ein Framework für Approximate Q-Learning mit linearer Funktionsapproximation im beliebten [Snake-Spiel](https://en.wikipedia.org/wiki/Snake_(video_game_genre)). Lies die Anweisungen sorgfältig durch und vervollständige die unvollständigen Funktionen. Anschließend kannst du das Trainingsverfahren ausführen und deine Implementierung überprüfen.

Detaillierte Anweisungen:

1. Im Skript *"snake.py"* findest du die Definition der *"Snake"*-Klasse.
2. **Falls du Google Colab verwendest: Lade die Datei *snake.py* hoch, indem du sie per Drag & Drop links in die Ordnerstruktur ziehst.**
3. Wichtige Begriffe und Spielregeln:
    - **state (Zustand)**: Ein Zustand wird durch die Positionen der Snake-Elemente auf dem Bildschirm (Breite: 30, Höhe: 20) und die Position des Essens bestimmt. Zu Beginn einer Episode wird das Spiel mit einer Schlange der Länge 1 und einer zufälligen Position für das Essen initialisiert. Jedes Mal, wenn sich der Kopf der Schlange an der gleichen Position wie das Essen befindet, wird die Länge der Schlange um 1 erhöht, der Punktestand steigt um 1, und ein neues Essen erscheint an einer zufälligen Position. <br>
    - **action (Aktion)**: In jedem Zustand gibt es 4 mögliche Aktionen, nämlich die Richtung, in die sich der Schlangenkopf um 1 Feld bewegen kann (oben, rechts, unten, links).  <br>
    - **termination (Beendigung)**: Wenn sich die Schlange in ihren eigenen Körper bewegt oder den Bildschirm verlässt, endet das Spiel. <br>
    - **features (Merkmale)**: Ähnlich zu den Tetris-Merkmalen aus den Vorlesungsfolien (Folie 28) extrahieren wir handgefertigte Merkmale des aktuellen Zustands, auf deren Basis die Q-Werte berechnet werden können. Die insgesamt 6 Merkmale sind:
        - `pos_x`: Befindet sich das Essen rechts von der Schlange, ist dieser Wert 1. Befindet es sich links von der Schlange, ist der Wert -1. Hat es die gleiche x-Koordinate wie die Schlange, ist der Wert 0.
        - `pos_y`: Befindet sich das Essen unterhalb der Schlange, ist dieser Wert 1. Befindet es sich oberhalb der Schlange, ist der Wert -1. Hat es die gleiche y-Koordinate wie die Schlange, ist der Wert 0.
        - `surrounding`: Enthält 4 Einträge. Jeder Eintrag repräsentiert das Feld direkt über, rechts, unter oder links vom Schlangenkopf. Ist dieses Feld belegt (entweder durch die Schlange oder eine Wand), beträgt dieser Eintrag 1, ansonsten 0.

    <br>

4. **Darstellung der Q-Werte**:
    Für einen gegebenen Zustand $s$ haben wir immer 4 Q-Werte $Q(s, a_i) \in \mathbb{R}$, einen für jede der vier Aktionen $a_1, a_2, a_3, a_4$. Wir stellen sie als vierdimensionalen Vektor dar und berechnen sie mit der Formel:

    $$ Q(s) = (Q(s, a_1), ..., Q(s,a_4)) = \phi(s)^T \beta \in \mathbb{R}^4.$$

Hierbei ist $\phi(s) \in \mathbb{R}^6$ die Merkmalsrepräsentation des Zustands. Unsere lernbaren Gewichte $\beta \in \mathbb{R}^{6 \times 4}$ sind eine Matrix, wobei jede Zeile für die Berechnung des Q-Werts einer der 4 Aktionen verantwortlich ist.

4. Weitere funktionsspezifische Anweisungen befinden sich über den unvollständigen Funktionen.


```python
##### DO NOT CHANGE #####
# Infer if we are on google colab
try:
    from google.colab.patches import cv2_imshow
    colab_mode = True
except ModuleNotFoundError:
    colab_mode = False

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# Include some python packages and snake env
import os
import pygame
import random
from snake import Snake
import time
import numpy as np
if colab_mode:
    import cv2
    from google.colab.patches import cv2_imshow
    from google.colab import output as colab_output

grading_env_var = os.getenv('NBGRADER_EXECUTION', None)
is_grading = (grading_env_var == "autograde" or grading_env_var == "validate")

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# initialization

# fixed seed for deterministic behavior:
np.random.seed(100)
random.seed(100)

# dimensions
feature_size = 6
action_size = 4


# learnable weight initialization
# these are the global weights we update during the learning
betas = np.random.rand(feature_size, action_size)

# number of transitions the agent performs Q learning
len_epoch = 20000
# hyperparameters
temp = 100.0
lr = 0.2
gamma = 0.5

# action translator. The snake environment expects string keywords, but for the RL agent it is simpler to output
# indices between 0 and 3.
actions = {
    0: "left",
    1: "up",
    2: "right",
    3: "down"
}

##### DO NOT CHANGE #####
```

### 2 a) Value Function Approximation (1 Punkt) <br>
Bitte vervollständige die Funktion **"approximate_q_values"** unten. Diese Funktion soll eine **"lineare Funktionsapproximation"** durchführen, um die 4 Q-Werte $Q(s, a)$ für einen Zustand zu berechnen.<br>

Hinweis:
- Du solltest die Parameter **"betas"** sowie einige nützliche numpy-Funktionen wie **"np.matmul"** oder **"np.dot"** verwenden.<br> Beachte, dass die lineare Funktionsapproximation durch ein Skalarprodukt implementiert werden kann.
- Die korrekte Formel findest du in den oberen Anweisungen unter bulletpoint 4.
- Du kannst die Korrektheit deiner Funktion mit der Testzelle unten überprüfen.


<!-- grade_id: cell-52c0d168df750e08, points: 0, grade: false, solution: true -->
```python
# TODO: PLEASE FINISH THE FUNCTION IN THIS BLOCK
def approximate_q_values(phi_s):
    """
    phi_s: Feature vector with shape (feature_size,). Extracted features from a state s.

    return: Vector containing the 4 Q-values Q(s,a) for every possible action a.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return approximated_q_value
```

<!-- grade_id: Ex2a, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: Ex2a - possible points: 1

### Testing your code (If your implementation is right, the following assertion
### on randomly generated inputs should hold.)

phi_s_example = np.array([0.7, 1.4, 1.1, 0.5, 0.2, 0.9])
assert np.allclose(approximate_q_values(phi_s_example), np.array([1.18090161, 1.93226149, 3.22564332, 2.82428688]))
print("Test Passed!")

##### DO NOT CHANGE #####
```

### 2 b) Exploration vs. Exploitation (3 Punkte) <br>
Bitte vervollständige die Funktionen **"select_action"** und **"softmax"** unten. Die Funktion **"select_action"** wählt eine Aktion (gibt ihren Index zurück) gemäß der **"Soft-Max-Policy"**-Explorationsstrategie aus (siehe Folie 18 im Kapitel Reinforcement Learning). Erinnere dich daran, dass diese Explorationsstrategie einen **"Temperatur"**-Wert besitzt (in unserem Notebook `temp` genannt), der die Q-Werte skaliert. Eine höhere Temperatur führt zu ähnlicheren Eingaben für die Soft-Max-Funktion, was eine gleichmäßigere Verteilung der Auswahlwahrscheinlichkeiten zur Folge hat. Eine Temperatur nahe 0 erhöht die Schärfe der Wahrscheinlichkeitsverteilung, sodass fast immer die Aktion mit dem höchsten Q-Wert gewählt wird. <br>

- Beginne mit der Implementierung der **"softmax"**-Funktion:
$$
f: \mathbb{R}^n \rightarrow \mathbb{R}^n
$$
$$
f(x)_i = \frac {\exp(x_i)}{\sum_j \exp(x_j)}.
$$
Beachte, dass sowohl die Eingabe als auch die Ausgabe Vektoren mit der gleichen Form sind. In dieser Funktion gibt es keine Temperatur, stattdessen skalieren wir die Eingaben zu dieser Funktion.

- Implementiere anschließend die **"select_action"**-Funktion. Sie hat zwei Argumente: Die aktuelle Merkmalsrepräsentation `phi_s` sowie einen booleschen Wert `sample`. Falls `sample` auf `True` gesetzt ist, soll die Aktion basierend auf der Soft-Max-Policy ausgewählt werden (denke daran, hier die Werte mit der Temperatur zu skalieren). Andernfalls soll die deterministische Aktion mit dem höchsten Q-Wert zurückgegeben werden (wobei die Temperatur ignoriert wird).

Hinweise:
- Die Softmax-Funktion kann für eine bessere numerische Stabilität anders implementiert werden als in der oben angegebenen Formel. In dieser Übung kannst du jedoch die einfachere Implementierung gemäß der Formel verwenden.
- Du kannst die Funktion **"np.argmax"** nutzen, um den Index des maximalen Wertes in einem `np.array` zu erhalten.
- Verwende die Softmax-Funktion, um die Auswahlwahrscheinlichkeiten zu berechnen. Denke daran, die Q-Werte mit der Temperatur zu skalieren. Diese ist als globale Variable `temp` definiert.
- Sobald du die Auswahlwahrscheinlichkeiten hast, kannst du die Funktion `np.random.multinomial(...)` für das Sampling nutzen. Sieh dir die Dokumentation hier an: [np.random.multinomial](https://numpy.org/doc/stable/reference/random/generated/numpy.random.multinomial.html).
- Stelle sicher, dass der Rückgabewert ein Integer ist und kein `np.array` mit Länge 4!
- Du kannst die Korrektheit deiner Softmax-Implementierung mit dem Test in der Zelle unten überprüfen.

<!-- grade_id: cell-f34d8b77fe0a975a, points: 0, grade: false, solution: true -->
```python
def softmax(x):
    """
    x: 1-dimensional numpy array.
    returns: softmax(x), 1-dimensional numpy.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return softmax_value

def select_action(phi_s, sample=True):
    """
    phi_s: Feature vector with shape (feature_size,). Extracted features from a state s.
    sample: Boolean flag. If true, sample an action based on the Soft-Max policy, otherwise return the action with
            the highest Q-value.

    return: Integer action index a_idx (in [0,1,2,3]) of the selected action.
    """
    qs = approximate_q_values(phi_s)
    if sample:
    # YOUR CODE HERE
    raise NotImplementedError()
    else:
    # YOUR CODE HERE
    raise NotImplementedError()

    return a_idx
```

<!-- grade_id: Ex2b_softmax, points: 1, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: Ex2b_softmax - possible points: 1

### Testing your code (If your implementation is right, the following assertion
### on randomly generated inputs should hold!!)
# example input
x = np.array([1.0, 2.0, 3.0, 4.0])
output = softmax(x)
assert np.allclose(output, np.array([0.0320586,  0.08714432, 0.23688282, 0.64391426]))
print("Test passed!")

##### DO NOT CHANGE #####
```

<!-- grade_id: ex2b_select_action, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: ex2b_select_action - possible points: 2

# Testing your code (If your implementation is right, the following assertion should hold!!)
np.random.seed(2201)
random.seed(2201)
betas = np.random.rand(feature_size, action_size)

phi_s_example = np.array([0.5, 1.0, 1.1, 0.2, 0.3, 0.4])
phi_s_example_2 = np.array([0.7, 1.4, 0.8, 0.3, 0.1, 0.7])
assert select_action(phi_s_example) == 0
print("Test Passed!")

##### DO NOT CHANGE #####
```

### 2 c) Berechnung des Temporal Difference Fehlers (2 Punkte) <br>
Bitte vervollständige die Funktion **"compute_delta"** unten. Diese Funktion berechnet den Temporal Difference Error $\delta$, wie er im Algorithmus auf Folie 31 des RL-Kapitels angegeben ist. Allerdings müssen wir auch terminale Zustände berücksichtigen, was auf der Folie nicht dargestellt ist. Falls der Zustand $s$ terminal ist, berechnen wir $\delta$ mit der Formel:

$$
\delta = r(s,a) - Q(s,a).
$$
<br>

Hinweise:
- Verwende deine implementierte **"approximate_q_values()"**-Funktion.
- Du kannst die Funktion **"np.max"** nutzen, um den maximalen Wert eines `np.array` zu erhalten.
- Verwende die global initialisierte Variable **"gamma"** als Diskontfaktor.

<!-- grade_id: cell-8832f0e67cac7236, points: 0, grade: false, solution: true -->
```python
# TODO: PLEASE FINISH THE FUNCTION IN THIS BLOCK
def compute_delta(phi_s, a_idx, r, phi_new_s, is_terminal):
    """
    phi_s: Feature vector of state s with shape (feature_size,).
    a_idx: action index a_idx (in [0,1,2,3]) of the action selected in state s.
    r: reward r(s,a) of the action with index a_idx in state s (float).
    phi_new_s: Feature vector of state new_s with shape (feature_size,).
               new_s is the state after selection action a_idx in state s.
    is_terminal: boolean which indicates if s is a terminal state.

    return: td_error delta of type float
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return delta
```

<!-- grade_id: ex_2c, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: ex_2c - possible points: 2

np.random.seed(60)
random.seed(60)
betas = np.random.rand(feature_size, action_size)
phi_s = np.array([0.5, 1.0, 1.1, 0.2, 0.3, 0.4])
a_idx = 2
r = 4
phi_new_s = np.array([0.7, 1.4, 0.8, 0.3, 0.1, 0.7])
is_terminal = False
assert np.allclose(compute_delta(phi_s, a_idx, r, phi_new_s, is_terminal), 3.2786106153158)
is_terminal = True
assert np.allclose(compute_delta(phi_s, a_idx, r, phi_new_s, is_terminal), 2.287845163104087)
print("Test Passed!")

##### DO NOT CHANGE #####
```

### 2 d) Berechnung der Ableitung des Q-Werts bezüglich Beta (2 Punkte) <br>
Bitte vervollständige die Funktion **"compute_d_qsa_d_beta"** unten. Diese Funktion berechnet die Ableitung des Q-Werts bezüglich des globalen Parameters **"betas"**. Daher musst du zunächst verstehen, was die Ableitung $ \frac {\text{d}Q(s,a)}{\text{d} \beta} $ genau ist, bevor du sie implementierst. Einige allgemeine Hinweise:
<br>


Wenn wir die Spalten von $\beta$ definieren als
$$
\beta  = (\beta_1, \beta_2, \beta_3, \beta_4)
$$
erhalten wir 4 Vektoren $\beta_i \in \mathbb{R}^6$.

Falls die Aktion $a$ den Index $i_a$ hat, ergibt sich:
$$
Q(s, a) = \phi(s)^T\beta_{i_a}.
$$

Die Ableitung einer Matrix kann als die Ableitung bezüglich jeder Spaltenvektor betrachtet werden, horizontal gestapelt:

$$
\frac {\text{d}Q(s,a)}{\text{d} \beta} = \left(\frac {\text{d}Q(s,a)}{\text{d} \beta_{i,j}}\right)_{i,j} = \left(\frac {\text{d}Q(s,a)}{\text{d} \beta_{1}}, ... \frac {\text{d}Q(s,a)}{\text{d} \beta_{4}} \right)
$$

Für $j \ne i_a$ hängt der Q-Wert nicht von $\beta_j$ ab, daher ist die Ableitung $0 \in \mathbb{R}^6$. Für $i_a$ ist die Ableitung $\phi(s) \in \mathbb{R}^6$.

**Daher ist die vollständige Ableitung
$$
\frac {\text{d}Q(s,a)}{\text{d} \beta}
$$
eine Matrix, die überall Nullen enthält, außer in der Spalte $i_a$, wo die Werte von $\phi(s)$ stehen.**


<!-- grade_id: cell-4d5d1d8a40d6854d, points: 0, grade: false, solution: true -->
```python
# TODO: PLEASE FINISH THE FUNCTION IN THIS BLOCK
def compute_d_qsa_d_beta(phi_s, a_idx):
    """
    phi_s: Feature vector of state s with shape (feature_size,).
    a_idx: action index a_idx (in [0,1,2,3]) of the action selected in state s.

    return: Derivative of the q_value Q(s, a) wrt. betas. It has shape (feature_size, action_size)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return d_qsa_d_beta
```

<!-- grade_id: ex_2d, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: ex_2d - possible points: 2

phi_s = np.array([0.5, 1.0, 1.1, 0.2, 0.3, 0.4])
a_idx = 0

d_qsa = compute_d_qsa_d_beta(phi_s, a_idx)
assert np.allclose(d_qsa, np.array([[0.5, 0. , 0. , 0. ],
 [1.  , 0. , 0.,  0. ],
 [1.1 ,0.  ,0.  ,0. ],
 [0.2 ,0.  ,0.  ,0. ],
 [0.3 ,0.  ,0.  ,0. ],
 [0.4 ,0.  ,0.  ,0. ]]))
print("Test Passed!")

##### DO NOT CHANGE #####
```

### 2 e) Gradientabstieg (2 Punkte) <br>
Bitte vervollständige die Funktion **"update_betas"** unten. Diese Funktion aktualisiert die Parameter **"betas"**, die in der Q-Wert-Funktionsapproximation verwendet werden. Die Formel dazu findest du auf Folie 31 (Update des Parametervektors).


Hinweise:
- Verwende deine Implementierung des TD-Fehlers $\delta$ sowie der Ableitung der Q-Wert-Funktion $\frac {\text{d}Q(s,a)}{\text{d} \beta}$.
- Gib den neuen Wert von **betas** zurück.


<!-- grade_id: cell-4e07fac9f25ea6c3, points: 0, grade: false, solution: true -->
```python
# TODO: PLEASE FINISH THE FUNCTION IN THIS BLOCK
def update_betas(phi_s, a_idx, r, phi_new_s, is_terminal):
    """
    phi_s: Feature vector of state s with shape (feature_size,).
    a_idx: action index a_idx (in [0,1,2,3]) of the action selected in state s.
    r: reward r(s,a) of the action with index a_idx in state s (float).
    phi_new_s: Feature vector of state new_s with shape (feature_size,).
               new_s is the state after selection action a_idx in state s.
    is_terminal: boolean which indicates if s is a terminal state.

    return: new value of betas (same shape as betas)
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    return new_betas
```

<!-- grade_id: ex2e, points: 2, grade: true, solution: false -->
```python
##### DO NOT CHANGE #####
# ID: ex2e - possible points: 2

# Hidden tests..

##### DO NOT CHANGE #####
```

Helper Functions for evaluating and playing the game:

```python
##### DO NOT CHANGE #####
# helper function to test the current policy
def test_policy(num_games=5):
    av_score = 0
    g = 0
    snake = Snake(FRAMESPEED=50000)
    while g < num_games:
        s = snake.get_feature_representation()
        a_idx = select_action(s, sample=False)
        a = actions[a_idx]
        is_terminal = snake.step(a)
        if is_terminal:
            av_score += snake.last_score
            g += 1
    pygame.quit()

    av_score /= num_games
    print(f"Average score: {av_score}")

##### DO NOT CHANGE #####
```

```python
##### DO NOT CHANGE #####
# helper function to simulate one game with normal speed
def play_single_game(framespeed=20):
    screen = pygame.display.set_mode((400, 300))
    snake = Snake(FRAMESPEED=framespeed)
    while True:
        s = snake.get_feature_representation()
        a_idx = select_action(s, sample=False)
        is_terminal = snake.step(actions[a_idx], init_new_game_after_terminal=False)
        if is_terminal:
            print(f"Total Score: {snake.last_score}")
            break
        if colab_mode:
            # Do a sketchy external visualization
            os.environ["SDL_VIDEODRIVER"]="dummy"  
            #convert image so it can be displayed in OpenCV
            view = pygame.surfarray.array3d(screen)
            #  convert from (width, height, channel) to (height, width, channel)
            view = view.transpose([1, 0, 2])

            #  convert from rgb to bgr
            img_bgr = cv2.cvtColor(view, cv2.COLOR_RGB2BGR)

            #Display image, clear cell every 0.1 seconds
            cv2_imshow(img_bgr)
            time.sleep(0.1)
            colab_output.clear()

##### DO NOT CHANGE #####
```

### Der vollständige Q-Learning-Algorithmus <br>
Nun sind alle einzelnen Module des Q-Learning vorbereitet. In der Zelle unten findest du den vollständigen Q-Learning-Algorithmus. Du solltest eine durchschnittliche Punktzahl von über 35 erreichen!


```python
# import the global beta values
global betas
import os

if not is_grading:

    snake = Snake(FRAMESPEED=50000)

    for i in range(len_epoch):
        # 1. get the feature representation of the current state by calling snake.get_feature_representation()
        phi_s = snake.get_feature_representation()
        # 2. Select the action using the epsilon-greedy exploration strategy
        a_idx = select_action(phi_s, sample=True)
        # 3. Get the action string which is needed for the snake environment. (already implemented)
        a = actions[a_idx]
        # 4. Ask for the current reward (already impelmented)
        r = snake.get_reward(a)
        # 5. Perform one step. This returns the boolean is_terminal if the snake died during this step.
        # (already implemented)
        is_terminal = snake.step(a)
        # 6. Get the feature representation of the updated state by calling snake.get_feature_representation()
        phi_new_s = snake.get_feature_representation()
        # 7. Update betas by calling the update_betas(...) method (already implemented)
        betas = update_betas(phi_s, a_idx, r, phi_new_s, is_terminal)

        # To see how well the current policy works we test it every 5000 updates
        if i % 5000 == 0:
            pygame.quit()
            test_policy()
            snake = Snake(FRAMESPEED=50000)

        # update hyperparameters
        temp *= 0.999
        temp = max(temp, 0.1)
        #print(epsilon)
        lr *= 0.9999

    test_policy()
    pygame.quit()
```

```python
# Test your learned strategy. You can adjust the speed of the animation with the framespeed argument. We used a cheesy way to make pygame work inside google colab, for a more smooth and faster animation, please run the code using your own jupyter notebook.
if not is_grading:
    play_single_game(framespeed=50)
```

```python
# if you want to close the pygame window (only relevant if you are outside google colab)
pygame.quit()
```

## Zusammenfassung
Das war die letzte Übung der Vorlesung GKI. Die hier vorgestellte Reinforcement Learning Methode wird nicht gut skalieren zu größeren Problemen, da Target Networks und Replay Buffers nicht implementiert wurden. Auch die Features sind nicht aussagekräftig genug, um globale Problemkonstellationen zu erkennen. Hoffentlich habt ihr aber einen Eindruck bekommen, wie ein RL-Agent aussehen kann.
Viel Erfolg bei der Klausur!